<!-- BEGIN QUESTION -->

## Exercise 03.1

Compare the computed values of 

$$ 
d_0 = a \cdot b + a  \cdot  c
$$

and 

$$ 
d_1 = a \cdot (b + c)
$$

when $a = 100$, $b = 0.1$ and $c = 0.2$. Store $d_{0}$ in the variable `d0` and $d_{1}$ in the variable `d1`.


Try checking for equality, e.g. `print(d0 == d1)`. 

<!-- grade_id: cell-447110c225124fc4, points: 0, grade: false, solution: true -->
```python
a = 100
b = 0.1
c = 0.2

d0 = a*b + a*c
print(d0)

d1 = a * (b+c)
print(d1)

print(d0 == d1)
```

Output:
```
30.0
30.000000000000004
False
```

<!-- grade_id: cell-91f2d89d88ea86c0, points: 0, grade: true, solution: false -->
```python
## tests ##
assert d0 == 30.0
assert d1 != 30.0
assert d0 != d1
```

<!-- END QUESTION -->

## Exercise 03.2

For the polynomial 
\begin{align}
f(x, y) &= (x + y)^{6} 
\\
&=  x^6 + 6x^{5}y + 15x^{4}y^{2} + 20x^{3}y^{3} + 15x^{2}y^{4} + 6xy^{5} + y^{6}
\end{align}
compute $f$ using: (i) the compact form $(x + y)^{6}$; and (ii) the expanded form for:

(a) $x = 10$ and $y = 10.1$

(b) $x = 10$ and $y = -10.1$

and compare the number of significant digits for which the answers are the same.
Store the answer for the compact version using the variable `f0`, and using the variable `f1` for the expanded version.

For case (b), compare the computed and analytical solutions and consider the relative error.
Which approach would you recommend for computing this expression?

#### (a) $x = 10$ and $y = 10.1$

```python
from math import factorial

#function that computes the expanded value of (x+y)^n 
def expanded(x,y,n):
    value = 0
    for r in range(0, n+1):
        value += (factorial(n)/(factorial(r) * factorial(n-r))) * (x**r) * (y**(n-r))
    
    return value

```

<!-- grade_id: cell-e8428e3de87d5500, points: 0, grade: false, solution: true -->
```python
x = 10.0
y = 10.1

f0 = (x+y) ** 6
f1 = expanded(x,y,6)

```

<!-- grade_id: cell-55d184514edf3a67, points: 0, grade: true, solution: false -->
```python
## tests ##
import math
assert math.isclose(f0, 65944160.60120103, rel_tol=1e-10)
assert math.isclose(f1, 65944160.601201, rel_tol=1e-10)
```

#### (b) $x = 10$ and $y = -10.1$

<!-- grade_id: cell-1561d20281ff67f6, points: 0, grade: false, solution: true -->
```python
x = 10.0
y = -10.1

f0 = (x+y) ** 6
f1 = expanded(x,y,6)

print(f0,f1)

```

Output:
```
9.999999999999788e-07 9.96515154838562e-07
```

```python
## tests ##
import math
assert math.isclose(f0, 1.0e-6, rel_tol=1e-10)
assert math.isclose(f1, 1.0e-6, rel_tol=1e-2)
```

## Exercise 03.3

Consider the expression

$$
f = \frac{1}{\sqrt{x^2 - 1} - x}
$$

When $x$ is very large, the denominator approaches zero, which can cause problems.

Try rephrasing the problem and eliminating the fraction by multiplying the numerator and denominator by $\sqrt{x^2 - 1} + x$ and evaluate the two versions of the expression when:

(a) $x = 1 \times 10^{7}$

(b) $x = 1 \times 10^{9}$ (You may get a Python error for this case. Why?)

#### Workings for rearranged form:

$$
\begin{align}
f &= \frac{1}{\sqrt{x^2-1} - x} \\
&= \frac{\sqrt{x^2-1} + x}{(\sqrt{x^2-1} - x)(\sqrt{x^2-1} + x)} \\ \\
&= \frac{\sqrt{x^2-1} + x}{x^2 -1 -x^2} \\
& = -\left(\sqrt{x^2-1} + x \right) \\
\end{align}
$$

#### (a) $x = 1 \times 10^{7}$

<!-- grade_id: cell-149ac8791e7ff020, points: 0, grade: true, solution: true -->
```python
def expression1(x):
    return 1/((x**2 - 1)**0.5 - x)

def expression2(x):
    return -1 * (((x**2 - 1)**0.5) + x)

```

```python
x= 1e7
print("original Form:", expression1(x))
print("new form:", expression2(x))
```

Output:
```
original Form: -19884107.85185185
new form: -19999999.999999948
```

#### (b) $x = 1 \times 10^{9}$

<!-- grade_id: cell-4ef986ef1f5838ba, points: 0, grade: true, solution: true -->
```python
x=1e9
#print("original Form:", expression1(x))
print("new form:", expression2(x))
```

Output:
```
new form: -2000000000.0
```

Python error because when $x=1\times 10^{-9}, \sqrt{x^2 - 1} \approx x$ and therefore $\sqrt{x^2 - 1} - x \approx 0$. Since the computer can only store 64 bits, including the sign, which results in 17 sig. fig.'s the calculated difernce between the computer calculated values of $\sqrt{x^2 - 1}$ and $x$ is 0, thus giving us a division by 0 error when running the original form.