## Exercise 09.1 (checking data validity)

The Fibonacci series is valid only for $n \ge 0$. Add to the Fibonacci function in this notebook a check that raises an exception if $n < 0$. Try some invalid data cases to check that an exception is raised.

*Optional:* Use `pytest` to test that an exception *is* raised for some $n < 0$ cases.

<!-- grade_id: cell-982c3611466532a2, points: 0, grade: false, solution: true -->
```python
def f(n, count=1, prev1=1, prev2 = 0): 
    "Compute the nth Fibonacci number using recursion"
    if n<0:
        raise ValueError("n must be greater than 0")
    
    elif n==0:
        return prev2
    
    elif n > count:
        #print(prev1, prev2)
        prev1, prev2 = (prev1 + prev2), (prev1)
        #print(prev1, prev2)
        count+=1
        return f(n, count, prev1, prev2)
    
    else:
        return(prev1)
    
```

<!-- grade_id: cell-9b97ea3e85d77f4a, points: 0, grade: true, solution: false -->
```python
## tests ##

# Perform some tests    
assert f(0) == 0
assert f(1) == 1
assert f(2) == 1
assert f(3) == 2
assert f(10) == 55
assert f(15) == 610

# Check that ValueError is raised for n < 0
import pytest
with pytest.raises(ValueError):
    f(-1)
with pytest.raises(ValueError):
    f(-2)
```

## Exercise 09.2 (raising exceptions)

Modify your program from the bisection exercise in Activity 04 to raise an error if the maximum number of iterations is exceeded. Reduce the maximum allowed iterations to test that an exception is raised.

Add any other checks on the input data that you think are appropriate.

<!-- grade_id: cell-ca2c3ee62ad3b956, points: 0, grade: false, solution: true -->
```python
def my_f(x):
    """Evaluate polynomial function"""
    return x**5 / 10 + x**3 - 10 * x**2 + 4 * x + 7

def compute_root(f, x0, x1, tol, max_it):
    # Iterate until tolerance is met
    it = 0
    error = tol + 1.0
    
    if type(max_it) != int or max_it<0:
        raise ValueError("max_it must be a positive integer")
    
    if tol<0:
        raise ValueError("tol must be positive ")
    
    if (type(x0) != int and type(x0) != float) or (type(x1) != int and type(x1) != float):
        raise ValueError("x0 and x1 must be numbers")
    
    while error > tol:
        # Compute midpoint
        x_mid = (x0 + x1) / 2

        # Evaluate function at (i) left end-point and at (ii) midpoint
        f0 = f(x0)
        f_mid = f(x_mid)

        if f0*f_mid < 0:
            x1 = x_mid
        else:
            x0 = x_mid

        error = abs(f_mid)
        it += 1

        # Guard against an infinite loop
        if it > max_it:
            raise RuntimeError("Exceeded Max Iterations")

    return x_mid, f_mid, it
```

<!-- grade_id: cell-f6cf91ed5dd1736b, points: 0, grade: true, solution: false -->
```python
## tests ##

# Test with max_it = 30
x, f, num_it = compute_root(my_f, x0=0, x1=2, tol=1.0e-6, max_it=30)

# Test with max_it = 20
with pytest.raises(RuntimeError):
    x, f, num_it = compute_root(my_f, x0=0, x1=2, tol=1.0e-6, max_it=20)
```